# <div style="background-color:rgba(204, 229, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(0, 76, 153, 1);">PHYS 231 Lab #4 - Day 2</span></div>

### Due Monday, Oct. 24 at 2:00 pm $-$ upload your .zip submission to the PHYS 231 Canvas Gradebook.

***
## Lab Learning Outcomes:
By the end of the PHYS 231 Lab #4, students will be able to:
* <b><span style="color:rgba(0, 153, 76, 1);">Design, build, and test practical $LRC$ circuits.</span></b>
* <b><span style="color:rgba(0, 153, 76, 1);">Measure the frequency response of an underdamped $LRC$ circuit.</span></b>
* <b><span style="color:rgba(0, 153, 76, 1);">Perform a weighted nonlinear fits to sets of data and use the extracted best-fit parameters to determine quantities of interest.</span></b>

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 1 - Installing & Importing Packages</span></div>

## The 'InstallerCheck()' Function:

Run the 'InstallerCheck()' function contained within PHYS231.py by placing you're cursor inside the code cell below and then hitting 'Shift'+'Enter'.  If the function reports that some packages have been installed, run the cell with "InstallerCheck()" a second time.  It should report that all required packages are already installed.

```python
import PHYS231
PHYS231.InstallerCheck()
```

## Importing Packages/Modules:

Once the packages are installed, they can be imported using the notation:
```python
import packageName
```
where 'packageName' is the name of the package to be imported.  Execute the import statements below by placing your cursor within the code cell and hitting 'Shift' + 'Enter'.

```python
# Import required and commonly-used modules
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd
import uncertainties
from uncertainties import ufloat
    
# Initialize Otter
import otter
grader = otter.Notebook("PHYS 231 - Lab 4b.ipynb")
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 2 - Oscilloscope Measurement of the Resonance and Phase</span></div>

## NumPy Arrays

Data that you collect during the lab can be entered directly into a list using the following notation:

```python
x = [x0, x1, x2, x3, ..., xN]
y = [y0, y1, y2, y3, ..., yN]
```
These lists will contain $N+1$ elements. It's usually more convenient to enter data into a NumPy array rather than a list because you can do mathematical operations on arrays that you can't do on lists.  For example:

- squaring an array creates an new array with each of the elements from the original array squared
- multiplying two arrays of the same length creates a new array in which the each elements is the product of the corresponding elements from the original arrays
- ...

NumPy arrays are entered using the following syntax:

```python
x = np.array([x0, x1, x2, x3, ..., xN])
y = np.array([y0, y1, y2, y3, ..., yN])
```             

In the cell below, replace the "..." in NumPy array statements with your measurements of:

- frequency in kHz
- the amplitude of the signal from the function generator V0 in volts
- the uncertainty in V0 in volts
- the amplitude of the voltage across the resistor VR in volts
- the uncertainty in VR in volts
- the difference in time in $\rm \mu s$ between V0 and VR at, say, a zero crossing with positive slope (note that some of the time differences will be positive and some will be negative)
- the uncertainty in the time difference in $\rm \mu s$

After executing the cell, the lengths of each of the arrays you've entered will be printed.  Make sure that all of the lengths are the same.

```python
# Enter your data into the arrays below.
# Please do NOT change the variable names.
frequency = np.array([...]) # kHz
V0 = np.array([...]) # V
errV0 = np.array([...]) # V
VR = np.array([...]) # V
errVR = np.array([...]) # V
dt = np.array([...]) # us
errdt = np.array([...]) # us

# The statements below will print the lengths of the arrays.
print("Length of frequency array:", len(frequency))
print("Length of V0 array:", len(V0))
print("Length of errV0 array:", len(errV0))
print("Length of VR array:", len(VR))
print("Length of errVR array:", len(errVR))
print("Length of dt array:", len(dt))
print("Length of errdt array:", len(errdt))
```

## Pandas DateFrames

The Pandas package can be used to display your array data in a nicely-formated table.  The Pandas package creates "DataFrames" which are 2 dimensional data structures, like a 2 dimensional array, or a table with rows and columns.  The cell below creates and then displays a DataFrame using the arrays that you created in the cell above.

```python
# Create the data structure.
data = {
    "frequnecy (Hz)": frequency,
    "V0 (V)": V0,
    "errV0 (V)": errV0,
    "VR (V)": VR,
    "errVR (V)": errVR,
    "dt (us)": dt,
    "errdt (us)": errdt
        }

# An option to format the numbers in the DataFrame.
pd.set_option('display.float_format', lambda x: '%0.3g' % x)

#Display the data.
pd.DataFrame(data)
```

## $V_\mathrm{R}/V_0$ versus $\omega = 2\pi f$

In the cells below, calculate:

- $\omega = 2\pi f$.  For $\pi$, use np.pi.  
- $V_\mathrm{R}/V_0$
- the uncertainty in $V_\mathrm{R}/V_0$

We'll assume the relative uncertainty in the frequency is negligible compared the relative uncertainties of the other relevant parameters.

Recall propagation of errors.  If $y = f\left(x1, x2, ..., x_N\right)$ and $x_i\pm\Delta x_i$ for $i= 1...N$ are known, then:

\begin{align}
\Delta y = \sqrt{\left(\frac{\partial f}{\partial x_1}\Delta x_1\right)^2 + \left(\frac{\partial f}{\partial x_2}\Delta x_2\right)^2 + ... + \left(\frac{\partial f}{\partial x_N}\Delta x_N\right)^2}
\end{align}

Please do **NOT** change the variable names.

```python
# Calucate the angular frequency in units of krad/s.
omega = ...
```

```python
grader.check("omega")
```

```python
Ratio = ...
```

```python
grader.check("Ratio")
```

```python
# Note that, in Python, the square root can be called using np.sqrt(...).
errRatio = ...
```

```python
grader.check("Ratio_uncertainty")
```

### Plot the data

The $V_\mathrm{R}/V_0$ data can now be plotted as a function of angular frequency.  Execute the cell below to created a scatter plot with error bars.  For additional information about the 'Scatter' function, see **Appendix A** near the end of this notebook. 

```python
fig1 = PHYS231.Scatter(
    xData = omega,
    yData = Ratio,
    yErrors = errRatio,
    xlabel = "angular frequency",
    ylabel = "VR/V0",
    xUnits = "krad/s",
    yUnits = "",
)
```

## Nonlinear Fits

Recall from class that, for the series $LRC$ circuit:

\begin{equation}
\frac{V_\mathrm{R}}{V_0} = \frac{R}{\sqrt{R^2 + \left(\omega L - \dfrac{1}{\omega C}\right)^2}}
\end{equation}

If the substitutions $\gamma = R/L$ and $\omega_0 = 1/\sqrt{LC}$ are made, one can show that:

\begin{equation}
\frac{V_\mathrm{R}}{V_0} = \frac{1}{\sqrt{1 + \left(\dfrac{\omega}{\gamma}\right)^2\left[1 - \left(\dfrac{\omega_0}{\omega}\right)^2\right]^2}}
\end{equation}

The parameter $\omega_0$ represents the resonance frequency and $\gamma$ represents the width of the resonance.

In practice, the peak of your resonance will not reach a value of $V_\mathrm{R}/V_0 = 1$.  This is due to losses on the wire used to wind the inductors which have not been taken into account.  To capture the effect of these losses, we can add an overall scaling factor $A$ such that:

\begin{equation}
\frac{V_\mathrm{R}}{V_0} = \frac{A}{\sqrt{1 + \left(\dfrac{\omega}{\gamma}\right)^2\left[1 - \left(\dfrac{\omega_0}{\omega}\right)^2\right]^2}}
\end{equation}

You will fit this function to your data in order to extract the best-fit parameters $A\pm\Delta A$, $\omega_0\pm\Delta\omega_0$, and $\gamma\pm\Delta\gamma$.  

For nonlinear functions, it is often necessary to provide initial "guesses" for the parameters which are used by the fitting algorithm.  Use the plot generated above to determine reasonable initial guesses for the parameters $A$, $\omega_0$, and $\gamma$.

```python
# Enter your starting parameters into array below.  The order must be np.array([A, omega0, gamma]).
start = np.array([...])
```

<!-- BEGIN QUESTION -->

We'll use the 'Lorentz' contained in PHYS231.py to complete the nonlinear fit:
```python
A, Omega0, Gamma, errA, errOmega0, errGamma, fig = PHYS231.Lorentz(x, y, erry)
```
A weighted fit with a formatted plot can be called by using some of the other available arguments:
```python
A, Omega0, Gamma, errA, errOmega0, errGamma, fig = PHYS231.Lorentz(xData = x, yData = y, yErrors = erry, start = start, xlabel = 'time', ylabel = 'position', xUnits = 's', yUnits = 'm')
```
Execute the cell below to preform a weighted nonlinear fit to your data.

```python
A, Omega0, Gamma, errA, errOmega0, errGamma, fig2 = PHYS231.Lorentz(
    xData = omega, 
    yData = Ratio, 
    yErrors = errRatio, 
    start = start,
    xlabel = "angular frequency", 
    ylabel = 'VR/V0', 
    xUnits = 'krad/s', 
    yUnits = '')
```

<!-- END QUESTION -->

In your lab notebook, you should compare this experimentally-measured values of $\omega_0$ and $\gamma$ to the values that you expected based on the values of the components used in your circuit.  Recall that $\omega_0 = 1/\sqrt{LC}$ and $\gamma=R/L$.  Keep in mind that your expected values of $\omega_0$ and $\gamma$ will also have uncertainties based on the uncertainties in the values of $R\pm\Delta R$, $L\pm\Delta L$, and $C\pm\Delta C$.

## $\phi$ versus $\omega$

In the cells below, using $\omega$ and the measured time difference $dt$ to calculate the phase difference between $V_\mathrm{R}$ and $V_0$ (and its uncertainty).  Recall that $omega$ is in $\rm krad/s$ and $dt$ is in $\rm \mu s$.  You will need to convert to $\rm rad/s$ and $\rm s$ to get the correct values for $\phi$.

Please do **NOT** change the variable names.

```python
Phi = ...
```

```python
grader.check("phase")
```

```python
errPhi = ...
```

```python
grader.check("phase_uncertainty")
```

### Plot the data

The $\phi$ data can now be plotted as a function of angular frequency.  Execute the cell below to created a scatter plot with error bars.  For additional information about the 'Scatter' function, see **Appendix A** near the end of this notebook. 

```python
fig3 = PHYS231.Scatter(
    xData = omega,
    yData = Phi,
    yErrors = errPhi,
    xlabel = "angular frequency",
    ylabel = "phase",
    xUnits = "krad/s",
    yUnits = "rad",
)
```

## Nonlinear Fits

Recall from class that, for the series $LRC$ circuit:

\begin{equation}
\phi = \arctan\left(\dfrac{1}{\omega RC} - \omega\frac{L}{R}\right)
\end{equation}

Expressed in terms of $\gamma = R/L$ and $\omega_0 = 1/\sqrt{LC}$, the phase becomes:

\begin{equation}
\phi = \arctan\left\{ \frac{\omega}{\gamma}\left[\left(\frac{\omega_0}{\omega}\right)^2 - 1\right] \right\}
\end{equation}

You will fit this function to your data in order to extract the best-fit parameters $\omega_0\pm\Delta\omega_0$ and $\gamma\pm\Delta\gamma$.  

For nonlinear functions, it is often necessary to provide initial "guesses" for the parameters which are used by the fitting algorithm.

```python
# Enter your starting parameters into array below.  The order must be np.array([omega0, gamma]).
start = np.array([...])
```

<!-- BEGIN QUESTION -->

We'll use the 'Phase' contained in PHYS231.py to complete the nonlinear fit:
```python
Omega0, Gamma, errOmega0, errGamma, fig = PHYS231.Phase(x, y, erry)
```
A weighted fit with a formatted plot can be called by using some of the other available arguments:
```python
Omega0, Gamma, errOmega0, errGamma, fig = PHYS231.Phase(xData = x, yData = y, yErrors = erry, start = start, xlabel = 'time', ylabel = 'position', xUnits = 's', yUnits = 'm')
```
Execute the cell below to preform a weighted nonlinear fit to your data.

```python
Omega0, Gamma, errOmega0, errGamma, fig4 = PHYS231.Phase(
    xData = omega, 
    yData = Phi, 
    yErrors = errPhi, 
    start = start,
    xlabel = "angular frequency", 
    ylabel = 'phase', 
    xUnits = 'krad/s', 
    yUnits = 'rad')
```

<!-- END QUESTION -->

In your lab notebook, you should compare this experimentally-measured values of $\omega_0$ and $\gamma$ to the values that you expected based on the values of the components used in your circuit.  Recall that $\omega_0 = 1/\sqrt{LC}$ and $\gamma=R/L$.  Keep in mind that your expected values of $\omega_0$ and $\gamma$ will also have uncertainties based on the uncertainties in the values of $R\pm\Delta R$, $L\pm\Delta L$, and $C\pm\Delta C$.

<!-- BEGIN QUESTION -->

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 3 - Feedback and Submission</span></div>

***
**<span style="color:blue">Question 3.1:</span>**  

We welcome your feedback on the PHYS 231 labs!  Please feel free to include any comments you have about this lab in the cell below.  Your comments will be taken into consideration when revising/improving the PHYS 231 labs.  You can suggest improvements, point out anything that was unclear, comment on the strengths and weaknesses of the lab, ...

This question is optional and will have no impact on your lab grade.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

***
Once you've completed this notebook:
- Save your work.
- Run 'grader.export()' to generate a .zip file containing all of the materials that you will submit.
- Download the generated .zip file.
- Upload the .zip file to the PHYS 231 Canvas gradebook.

Here is a <a href = "https://cmps-people.ok.ubc.ca/jbobowsk/phys231/Python/images/Submission.gif">GIF</a> showing how these steps are completed.  Once your completed notebook has been uploaded to the Canvas gradebook, you're done!

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 4 - Playground (optional)</span></div>

Feel free to add as many cells as you like below and use them as a playground for further independent investigations.  These cells won't be graded, so feel free to use them in any way that you like.  

```python
# Here's an empty code cell that you can use.
```

```python
# Here's another empty code cell that you can use.
```

```python
# Here's yet another empty code cell that you can use.
# If you need more, you can add cells using the '+' icon in the menu bar at to the top of the screen.
```

### <div style="background-color:rgba(255, 255, 102, 0.5); text-align:left; padding:20px 0; margin-top:20px">$\quad$Appendix A &ndash; The Scatter Function...</div>

The function for generating scatter plots is called as follows:
```python
PHYS231.Scatter(xData, yData, yErrors = [], xlabel = 'x-axis', ylabel = 'y-axis', xUnits = '', yUnits = '', fill = False, show = True)
```
The 'xData' and 'yData' inputs are required, all other arguments are optional with default values set.  The function returns the a single output (the formatted plot):
```python
fig
```

The function will do a simple scatter plot if no 'yError' are included.  It will included error bars if 'yErrors' are passed to the function.  The 'fill' and 'show' arguments should generally be false.  

If the 'xData' list is empty, as in:
```python
xData = []
```
then the x-axis will be the trial number.  For example, is 'xData' is empty and 'yData' has 10 elements, the x-axis will span 1 to 10.

### Scatter Example Implmentation
The code block below shows an implementation of 'Scatter'.
```python
import PHYS231
theta = [10, 20, 30] # degrees
T = [2.02, 1.95, 2.13] # s
errT = [0.02, 0.02, 0.03] # s
fig = PHYS231.Scatter(theta, T, errT, 'initial angle' , 'period', 'degrees', 's')
```

If you're interested in generating your own plots with customized formatting, see the following Python-based plotting tutorial: https://cmps-people.ok.ubc.ca/jbobowsk/Python/html/Jupyter%20Basic%20Plots.html.
***

Last update: October 14, 2022

Tests for Ratio (points: 1)
```
>>> len(Ratio) != 0
True
```
```
>>> test = True
>>> for i in range(len(Ratio)):
...     if np.isclose(Ratio[i], VR[i]/V0[i], atol = 0, rtol = 0.03) != True:
...         test = False
>>> test
True
```


Tests for Ratio_uncertainty (points: 1)
```
>>> len(errRatio) != 0
True
```
```
>>> test = True
>>> for i in range(len(errRatio)):
...     if np.isclose(errRatio[i], np.sqrt((errVR[i]/V0[i])**2 + (errV0[i]*VR[i]/V0[i]**2)**2), atol = 0, rtol = 0.03) != True:
...         test = False
>>> test
True
```


Tests for omega (points: 1)
```
>>> len(omega) != 0
True
```
```
>>> test = True
>>> for i in range(len(omega)):
...     if np.isclose(omega[i], 2*np.pi*frequency[i], atol = 0, rtol = 0.03) != True:
...         test = False
>>> test
True
```


Tests for phase (points: 1)
```
>>> len(Phi) != 0
True
```
```
>>> test = True
>>> for i in range(len(Phi)):
...     if np.isclose(Phi[i], (omega[i]*1e3)*(dt[i]*1e-6), atol = 0, rtol = 0.03) != True:
...         test = False
>>> test
True
```


Tests for phase_uncertainty (points: 1)
```
>>> len(errPhi) != 0
True
```
```
>>> test = True
>>> for i in range(len(errPhi)):
...     if np.isclose(errPhi[i], (omega[i]*1e3)*(errdt[i]*1e-6), atol = 0, rtol = 0.03) != True:
...         test = False
>>> test
True
```
